CLOCK evolved in cnidaria to synchronize internal rhythms with diel environmental cues

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #2 (Public Review):

In this revised manuscript Aguillon and collaborators convincingly demonstrating that CLK is required for free-running behavioral rhythms under constant conditions in the Cnidarian Nematostella. The results also convincingly show that CLK impacts rhythmic gene expression in this organism. This original work thus demonstrates that CLK was recruited very early during animal evolution in the circadian clock mechanism to optimize behavior and gene expression with the time-of-day. The manuscript could still benefit from some improvements so that it is more accessible for a wide readership.

Recommendations for the authors:

Reviewer #2 (Recommendations For The Authors):

Aguillon and collaborators have deeply revised, and in the progress significantly improved the presentation of their interesting results with the first Cnidarian circadian gene mutant. Results are now very convincingly demonstrating that CLK is required for free-running behavioral rhythms under constant conditions. The results also now more convincingly show that CLK impact rhythmic gene expression, although interpretation of the transcriptomics data is not straightforward. I think there is still improvements that are needed to make the manuscript more accessible. We authors need to keep in mind that a broad audience will read their report, not just chronobiologists. I have listed below several issues that I think should be addressed, and some editing suggestions.

General comment to Editor and Reviewers:

We are genuinely grateful to both reviewers and editors about all the feedback which helped us to make the best of our data, to question our analysis to the point we redefined our approach and end up with a great article we are proud of it. Only the name of authors is visible on the article, and considering how much the reviewing system help to improve the research it seems almost unfair. As such, we thank all of you and really appreciate the new eLife system. Bravo all.

Abstract:

(1) Line 40" It should read "transcript levels" instead of "transcription". There is no measurement of transcription rates in this manuscript, only mRNA levels.

Modified accordingly.

(2) Line 41: the authors mention "constant light". Does this refer to previous work? Their data in Figure 4 were in constant darkness, not in LL.

Modified accordingly.

(3) Line 46 and throughout the manuscript, the allelic nomenclature is not standard. 1-/- seems to indicate there are two different alleles. Since the allele might not be a null, I would suggest simply using 1/1, or perhaps delta/delta since the mutation results in a truncates CLK.

NvClk1-/- became NvClkΔ/Δ. Except in the .xls supplementary table were the mutant kept the NvClk-/- nomenclature. It is not possible to replace only part of a word with a different font, here generating delta sign would require to do it one by one.

(4) The last sentence of the abstract needs to be rephrased, as it suggests that CLK evolved to maintain circadian rhythms under constant conditions. Constant conditions very rarely exist on Earth, and thus cannot be an evolutionary driving force. Different explanations have been proposed on why a self-sustained clock is the evolutionary solution to timekeeping, but the purpose of the clock and of clock genes is not to maintain oscillations in constant conditions. Actually, this sentence conflicts with the title.

Modified to: the Clock gene has evolved in cnidarians to sustain 24-hour rhythmic physiology and behavior in absence of diel environmental conditions. From my actual understanding, you are right, the purpose of clock gene is not to maintain oscillation in constant conditions (this is simply the result of the experiment), but to synchronize the physiology to the day/night rhythm, and surely to sustain 24h oscillations in case the environment challenges the perception of the diel cues. The DD or LL is just an artificial experimental design to reveal the endogenous time-keeping pacemaker.

Results:

(1) Line 148 and elsewhere in the MS: I would not use the word "lower" or "higher" to qualify acrophases. I would suggest advanced/delayed or earlier/later.

Modified accordingly.

(2) Line 157-9: The introductory sentence does not clearly present the rationale for the 6/6 experiments.

We modified the paragraph accordingly: The presence of a 24-hour rhythm of NvClkΔ/Δ polyps under LD conditions could be attributed to either a direct light-response or the partial functioning of the circadian clock due to the nature of the mutation….

(3) At the end of the behavior section, or perhaps at the end of each paragraph in this section, it would be helpful to have a summary of the results and more clearly explain their interpretation. The authors need to guide the readers, particularly non-chronobiologist, so that they can understand what the really neat data that were obtained mean. For example, what does it mean that the acrophase is different between mutant and wild-type, why are Clk mutants rhythmic under LD12/12 or 6/6, etc.

We added a conclusion sentence to help non-specialist to understand each result.

(4) Line 172 and elsewhere" "true rhythmic genes" sounds odd to me. Either they are, or they are not rhythmic.

Modified to “rhythmic genes.”

(5) Paragraph starting with line 184: I do not follow what is important about the number of genes per time cluster. What does it tell us, beyond the simple fact that less genes are rhythmic in the Clk mutants?

We rewrote the result paragraph to make it clearer why we performed this clustering analysis. This clustering analysis became Extended Data Fig.2 with modification of the figures (see my comments in your review about Figure 3).

(6) Line 197: The authors need to explain what they saw with circadian clock genes and their expression in CLk mutants. In some case, amplitude increased in LD. This surprising observation deserves some explanations. "Complex regulatory effect" is too vague.

We replaced the vague “complex regulatory effect” by a more thorough description of the figure 3.a.

(7) Line 198-203: Again, help the reader understand the significance of these observations.

We rewrote the paragraph to help the reader to better understand the significance of these observations.

Discussion:

(1) Line 236-40. Careful with the use of -/-, which implies that an allele is a null. The first CLk mutants in mammals and flies, which the authors refer to. were actually dominant negatives.

I went over the citations we used for this paragraph and this first mutation in fly dClkar is null, no dominant negative. Flies are still rhythmic in the dark. Unless there is an older mutation? However, you right the first mutation identified in mouse was a dominant-negative with loss of rhythmicity, while the gene deletion did not show any effect on the behavior, suggesting compensation by a paralog. I removed two references which were not relevant to the discussion.

(2) Line 265-268 are not very clear. Do the authors mean that the lack of overlap for non-cricadian pacemaker genes is because of different experimental conditions? What would be those differences? It is reassuring that the Leach/Reitzel study and the present share pacemaker genes as rhythmic, but it is also surprising that there is almost no overlap beyond these genes. How robust are those other rhythms compared to circadian clock genes?

We revised this paragraph and raised major points regarding the raising condition of our polyps between labs and their potential genetic differences which could explain these differences.

(3) Line 270. I am not sure "compensation" is the right word, since there is no overlap between the rhythmic genes in mutants under LD and wild-type under either LD or DD. Also, saying on line 273 that the transcriptional pattern is not fully reproduced is a rather striking understatement, given the absence of rhythm gene overlap

We rewrote the paragraph accordingly. We replaced by “alternative way to drive rhythmicity under LD condition”.

(4) Line 279. The authors mention the possibility of false positives. Based on the FDR, is there more rhythmic genes than by chance?

The possibility of false-positive is a risk to consider when you do not perform multiple-testing. We added within the results paragraph the number of rhythmic genes identified with BH.Q or p.adj. which both are the multiple testing for each algorithm (RAIN and JTK) we used.

(5) Line 279-82. The references to the Ray study is rather obscure. What is the point the authors are trying to make here?

Eventually, we removed the reference from this article and modify the paragraph of the discussion. Indeed, the discussion around the Ray study did not gave an interesting direction to discuss our results and analysis approach.

(6) Line 284: define BHQ and p.adj

Defined and referenced.

(7) The way Lines 283-288 are worded do not provide a good rationale for how transcriptional rhythms were analyzed. The idea to combine two different approaches (JTK and RAIN) to be selective with rhythmicity was great. The authors need however to justify these choices in a more convincing manner. The goal is to detect rhythmic genes in a reliable manner, irrespective of the number of rhythmic genes observed Also, explaining the choice of methodology belongs to the result section.

We explained our choice of methodology and moved it to the result section as suggested.

(8) Line 292-3. There are known mechanisms that explain how transcriptional time clusters are generated. In particular, the use of interlocked feedback loop with antiphase peaks of transcriptions is well documented. Actually, it seems to me the clustering shown in Fig 4 might hint at such a mechanism.

Indeed you are right the clustering shown in Fig 3 (former Fig 4) revealed such mechanism.

Figures:

Figure 2: Define relative amplitude

We added the definition of the relative amplitude within the results. If this is what you asked for?

Figure 3: Some of the cycles look odd (first row of graphs in panel C). Why would the first and last data point be so different in three of these graphs?

We decided to modify this figure as we realized it was not informative and not objective enough, as we selected among multiple patterns few “representatives”. In the new figure we combined the cluster analysis to the behavior. Thus, readers can now pick a cluster according to a specific behavior activity level (or ZT/CT) and reach in supp. Table 4 the “genes of potential interest”. However generally speaking this figure does not explain more the consequences of the mutation, so we moved it into the Extended data Fig.2

Figure4: define the color coding in the correlation panels (blue to red)

These values from -1 to 1 are the Pearson correlation values. Now indicated on the figure with the color coding legend.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The circadian clock enables anticipation of the day/night cycle in animals ranging from cnidarians to mammals. Circadian rhythms are generated through a transcription-translation feedback loop (TTFL or pacemaker) with CLOCK as a conserved positive factor in animals. However, the functional evolutionary origin and mechanism of action of CLOCK in basal animals are unknown. In the cnidarian Nematostella vectensis, pacemaker genes transcription including NvClk (the Clock ortholog) appears arrhythmic under constant light conditions, questioning the role of NvCLK. Utilizing CRISPR/Cas9, we generated a NvClk allele mutant (NvClk1), revealing circadian behavior loss in constant light conditions (LL and DD) while a 24-hour rhythm was maintained under light-dark condition (LD). Transcriptomics showed distinct rhythmic genes in wild-type (WT) genes in LD compared to DD. The LD NvClk1-/- showed comparable numbers of rhythmic genes, whereas they were greatly reduced in DD. Furthermore, the LD NvClk1-/- showed alterations of temporal pacemaker genes expression, affecting their potential interactions. Additionally, differential expression of non-rhythmic genes associated with cell division and neuronal differentiation was observed. These findings suggest that while the light-responsive pathway can partially compensate for circadian clock disruption, the Clock gene has evolved in cnidarians to maintain 24-hour rhythmic physiology and behavior in constant conditions.

Revised version: The circadian clock enables anticipation of the day/night cycle in animals ranging from cnidarians to mammals. Circadian rhythms are generated through a transcription-translation feedback loop (TTFL or pacemaker) with CLOCK as a conserved positive factor in animals. However, CLOCK’s functional evolutionary origin and mechanism of action in basal animals are unknown. In the cnidarian Nematostella vectensis, pacemaker gene transcript levels, including NvClk (the Clock ortholog), appear arrhythmic under constant darkness, questioning the role of NvCLK. Utilizing CRISPR/Cas9, we generated a NvClk allele mutant (NvClkΔ), revealing circadian behavior loss in constant light conditions (LL and DD) while a 24-hour rhythm was maintained under light-dark condition (LD). Transcriptomics showed distinct rhythmic genes in wild-type (WT) genes in LD compared to DD. The LD NvClkΔ/Δ showed comparable numbers of rhythmic genes, but were greatly reduced in DD. Furthermore, the LD NvClkΔ/Δ showed alterations in temporal pacemaker gene expression, affecting their potential interactions. Additionally, differential expression of non-rhythmic genes associated with cell division and neuronal differentiation was observed. These findings revealed that a light-responsive pathway can partially compensate for circadian clock disruption and that the Clock gene has evolved in cnidarians to synchronize rhythmic physiology and behavior to the diel rhythm of the earth’s biosphere.

Introduction

Earlier version: Throughout the history of life on Earth, organisms have had to adapt to a constantly changing environment, including the ∼24-hour daily rhythm of light/dark, driving the development of endogenous biological clocks. The circadian clock, which is entrained by external stimuli such as light, enables the organism to anticipate the onset of the light and dark phases and synchronize its physiology and behavior in harmony with the environment. This, in turn, enhances the organism’s fitness and survival1–3. From single-celled organisms to metazoans, circadian clocks have evolved multiple times, highlighting their importance to living organisms1,2. Despite the fundamental role circadian clocks play in regulating the rhythmicity of living organisms, their evolutionary origin and intricate molecular mechanisms remain ambiguous in early diverging animal lineages, such as cnidaria.

Revised version: Throughout the history of life on Earth, organisms have had to adapt to a constantly changing environment, including the ∼24-hour daily rhythm of light/dark, driving the development of endogenous biological clocks. The circadian clock, entrained by external stimuli such as light, enables the organism to anticipate the onset of the light and dark phases and synchronize its physiology and behavior in harmony with the environment. This, in turn, enhances the organism’s fitness and survival1–3. From single-celled organisms to metazoans, circadian clocks have evolved multiple times, highlighting their importance to living organisms1,2. Despite the fundamental role circadian clocks play in regulating the rhythmicity of living organisms, their evolutionary origin and intricate molecular mechanisms remain ambiguous in early diverging animal lineages, such as cnidaria.

Earlier version: Rhythmic phenomena, including calcification, reproduction, and diel behavior patterns, have been examined in cnidarian species3–8. While some of these patterns were found to be directly triggered by environmental stimuli, such as light, others persist in the absence of external cues, suggesting the presence of an internally generated and self-sustaining circadian clock6,9,10. At the molecular level, cnidarians possess homologs of putative core pacemaker genes found in bilaterians11–13. Several studies have shown that most of these genes display diel expression patterns under light/dark cycles. However, in contrast to most animals, their oscillation generally ceases without the presence of light cues14–16. Thus, how the core pacemaker genes orchestrate rhythmic gene expression and circadian behaviors in cnidarians remains unclear.

Revised version: Rhythmic phenomena, including calcification, reproduction, and diel behavior patterns, have been examined in cnidarian species3–8. While environmental stimuli, such as light, directly triggered some of these patterns, others persist in the absence of external cues, suggesting the presence of an internally generated and self-sustaining circadian clock6,9,10. At the molecular level, cnidarians possess homologs of putative core pacemaker genes found in bilaterians11–13. Several studies have shown that most of these genes display diel expression patterns under light/dark cycles. However, unlike most animals, their oscillation generally ceases without light cues14–16. Thus, how the core pacemaker genes orchestrate rhythmic gene expression and circadian behaviors in cnidarians remains unclear.

Earlier version: One of the most studied cnidarian species in the field of chronobiology is the estuarine sea anemone, Nematostella vectensis. Few studies have shown that in diel lighting, the locomotor behavior of Nematostella has a ∼24-h rhythm that is maintained under constant conditions suggesting it is regulated by an endogenous circadian clock6,17–19. In accordance, the Nematostella genome codes for conserved core pacemaker genes such as NvClk, NvCycle, and the cryptochromes NvCry1a and NvCry1b12,13. The proposed circadian clock model in Nematostella is composed of the positive transcription factors (bHLH-PAS family), NvCLK and NvCYCLE, that heterodimerize and upregulate light-dependent cryptochrome genes in the feedback loop, and NvPAR-bZIPs in the feed-forward loop, which repress the transcription of the positive elements11,12,19. More recently, the NvCLK-interacting pacemaker, NvCIPC was predicted to act as an additional repressor of the NvCLK:NvCYCLE dimer19. However, in contrast to the free-running oscillation demonstrated for Nematostella behavior6,17,19, transcriptional expression profiles of most candidate genes implicated in the pacemaker do not retain their oscillation period in the absence of light14,15,19.

Revised version: One of the most studied cnidarian species in the field of chronobiology is the estuarine sea anemone, Nematostella vectensis. Few studies have shown that in diel lighting, the locomotor behavior of Nematostella has a ∼24-h rhythm that is maintained under constant conditions, suggesting it is regulated by an endogenous circadian clock6,17–19. By this, the Nematostella genome codes for conserved core pacemaker genes such as NvClk, NvCycle, and the cryptochromes NvCry1a and NvCry1b12,13. The proposed circadian clock model in Nematostella is composed of the positive transcription factors (bHLH-PAS family), NvCLK and NvCYCLE, that heterodimerize and upregulate light-dependent cryptochrome genes in the feedback loop, and NvPAR-bZIPs in the feed-forward loop, which repress the transcription of the positive elements11,12,19. More recently, the NvCLK-interacting pacemaker, NvCIPC, was predicted to act as an additional repressor of the NvCLK:NvCYCLE dimer19. However, in contrast to the free-running oscillation demonstrated for Nematostella behavior6,17,19, transcriptional expression profiles of most candidate genes implicated in the pacemaker do not retain their oscillation period without light14,15,19.

Earlier version: Here, we employed the CRISPR/Cas9-mediated genome editing system to establish a NvClk mutant (NvClk1-/-) Nematostella. By combining behavioral monitoring and transcriptomic analysis, we aimed to elucidate the role of NvClk in regulating rhythmic locomotor activity and gene expression under varying light conditions. Our study revealed a robust light response pathway capable of compensation and a conserved function of CLOCK as a timekeeper in the absence of a light cue.

Revised version: We employed the CRISPR/Cas9-mediated genome editing system to establish a NvClk mutant (NvClkΔ/Δ) Nematostella. By combining behavioral monitoring and transcriptomic analysis, we aimed to elucidate the role of NvClk in regulating rhythmic locomotor activity and gene expression under varying light conditions. Our study revealed a robust light response pathway capable of compensation and a conserved function of CLOCK as a timekeeper without a light cue.

Results

Phylogenetic analysis and spatial expression pattern of NvClk

Earlier version: Phylogenetic analysis of NvCLK protein sequences revealed that NvClk is positioned within the cnidarian branch (Fig. 1a). It contains a basic helix–loop–helix (bHLH) DNA binding domain and two Per-Arnt-Sim (PAS) domains, similar to the protein structure found in other animals. PAS domains are crucial structural motifs involved in protein-protein interactions that drive the self-sustaining molecular mechanism underlying the circadian clock20,21.

Revised version: Phylogenetic analysis of NvCLK protein sequences positioned NvClk within the cnidarian branch (Fig. 1a). It contains a basic helix–loop–helix (bHLH) DNA binding domain and two Per-Arnt-Sim (PAS) domains, similar to the protein structure found in other animals. PAS domains are crucial structural motifs in protein-protein interactions that drive the self-sustaining molecular mechanism underlying the circadian clock20,21.

Earlier version: NvClk1-/- cannot maintain circadian behavior in non-diel light conditions.

Revised version: NvClkΔ/Δ cannot maintain circadian behavior in non-diel light conditions.

Earlier version: (a) Phylogenetic tree showing the evolutionary relationship of CLK orthologs across different animal species. (b) In situ hybridization of NvClk in the WT juvenile with scale bars representing 0.1mm. (c) Schematic representation of the NvClk gene in grey, with the open reading frame (ORF) in dark grey and the conserved protein domains bHLH (yellow) and PAS1 and PAS2 (dark red). The CRISPR generated NvClk1 allele has a +20nt insertion after the PAS1 domain, represented by a black arrowhead. NvCLK dimerizes via its three functional domain with NvCYCLE binding the CACGTG ebox to drive rhythmic transcription. (d-g-j-m). Normalized Movement (a.u), hourly binned over 72h, under different light conditions: 12h light:12h dark, continuous dark (Dark - Dark), continuous light (Light - Light), and 6h light:6h dark. The black line represents the WT and the red line represents the NvClk1-/- mutant. (e-h-k-n) Lomb-Scargle Periodograms for each corresponding light condition. The significant period value (p<0.01) is indicated for each genotype in the top left corner of each graph. (f-i-l-o) Phase detection (Cosinor) and genotype comparison of 24h rhythmic individuals. See number nrhythmic/ntotal on the x-axis indicating for each genotype the number of 24h-rhythmic animals over the total population.

Revised version: (a) Phylogenetic tree showing the evolutionary relationship of CLK orthologs across different animal species. (b) In situ hybridization of NvClk in the WT juvenile with scale bars representing 0.1mm. (c) Schematic representation of the NvClk gene in grey, with the open reading frame (ORF) in dark grey and the conserved protein domains bHLH (yellow) and PAS1 and PAS2 (dark red). The CRISPR-generated NvClkΔ allele has a +20nt insertion after the PAS1 domain, represented by a black arrowhead. NvCLK dimerizes via its three functional domains with NvCYCLE binding the CACGTG ebox to drive rhythmic transcription. (d-g-j-m). Normalized Movement (a.u), hourly binned over 72h, under different light conditions: 12h light:12h dark, continuous dark (Dark - Dark), continuous light (Light - Light), and 6h light:6h dark. The black line represents the WT, and the red line represents the NvClkΔ/Δ mutant. (e-h-k-n) Lomb-Scargle Periodograms for each corresponding light condition. The significant period value (p<0.01) is indicated for each genotype in the top left corner of each graph. (f-i-l-o) Phase detection (Cosinor) and genotype comparison of 24h rhythmic individuals. See the number rhythmic/total on the x-axis indicating the number of 24h-rhythmic animals over the total population for each genotype.

Earlier version: To localize NvClk expression at the polyp stage, in situ hybridization chain reaction (HCRv.3) was performed. Polyps were sampled at ZT10, i.e., peak expression of NvClk6,12,22. NvClk specific expression was observed throughout the animal tissue, and enriched expression was visible in the tentacle endodermis and mesenteries, while no signal was observed in the negative control (Fig.1b, Extended Data Fig.1a). This expression pattern resembled the expression observed at the larvae stage22. To date, functional manipulation of the NvClk gene has not been performed in basal animal lineages including cnidarians, and its function is unknown in cnidarians20,21 (Fig.1a).

Revised version: In situ hybridization chain reaction (HCRv.3) was performed to localize NvClk expression at the polyp stage. Polyps were sampled at ZT10, i.e., peak expression of NvClk6,12,22. NvClk mRNA expression was observed throughout the animal tissue, and enriched expression was visible in the tentacle endodermis and mesenteries. In contrast, no signal was observed in the negative control (Fig.1b, Extended Data Fig.1a). This expression pattern resembled the expression observed at the larvae stage22. To date, functional manipulation of the NvClk gene has not been performed in basal animal lineages including, cnidarians, and its function is unknown in cnidarians20,21 (Fig.1a).

Earlier version: Generation of NvClk1-/- Nematostella

Revised version: Generation of NvClkΔ/Δ Nematostella

Earlier version: To investigate the function of NvClk in Nematostella, we employed the CRISPR-Cas9 system to generate mutants. Based on existing knowledge from mouse and Drosophila models, we hypothesized that NvCLK:NvCYCLE dimer binds to the DNA motif CACGTG within the promoter of rhythmic target genes (Fig. 1c). Guide RNA (gRNA) was synthesized to target a region between the two PAS domains of the NvClk coding sequence (CDS). This gRNA, along with the Cas9 endonuclease, was microinjected into zygotes (Methods). Subsequently, F0 animals were outcrossed with wild type (WT), and the F1 progeny were raised to adulthood. Genotyping of F1 polyps identified 10 different mutated alleles, with six displaying a frame-shift mutation, including one with a 20 bp insertion (NvClk1), resulting in a premature stop codon (Extended Data Fig. 1b). The predicted 203 amino acid truncated protein lacked 459 amino acids, including one co-factor dimerization PAS domain (Fig. 1c, Extended Data Fig. 1b). To obtain homozygous NvClk1-/- polyps, we crossed heterozygous NvClk1 F1 animals. Genotyping of F2 polyps confirmed the expected 25% frequency of NvClk1-/- mutants. Subsequently, we intercrossed NvClk1-/- animals to obtain F3 NvClk1-/- polyps for use in subsequent experiments aimed at assessing the role of NvClk in regulating behavioral and genetic rhythms.

Revised version: To investigate the function of NvClk in Nematostella, we employed the CRISPR-Cas9 system to generate mutants. Based on existing knowledge from mouse and Drosophila models, we hypothesized that NvCLK:NvCYCLE dimer binds to the DNA motif CACGTG within the promoter of rhythmic target genes (Fig. 1c). Guide RNA (gRNA) was synthesized to target a region between the two PAS domains of the NvClk coding sequence (CDS). This gRNA and the Cas9 endonuclease were microinjected into zygotes (Methods). Subsequently, F0 animals were outcrossed with wild type (WT), and the F1 progeny were raised to adulthood. Genotyping of F1 polyps identified ten different mutated alleles, with six displaying a frame-shift mutation, including one with a 20 bp insertion (NvClkΔ), resulting in a premature stop codon (Extended Data Fig. 1b). The predicted 203 amino acid truncated protein lacked 459 amino acids, including one co-factor dimerization PAS domain (Fig. 1c, Extended Data Fig. 1b). To obtain homozygous NvClkΔ/Δ polyps, we crossed heterozygous NvClkΔ F1 animals. Genotyping of F2 polyps confirmed the expected 25% frequency of NvClkΔ/Δ mutants. Subsequently, we intercrossed NvClkΔ/Δ animals to obtain F3 NvClkΔ/Δ polyps for use in subsequent experiments aimed at assessing the role of NvClk in regulating behavioral and genetic rhythms.

NvClk is necessary to maintain circadian behavior in constant conditions

Earlier version: To assess the impact of the NvClk1 mutation on circadian rhythm, we monitored locomotor behavior of WT and NvClk1-/- polyps under different light conditions (Supp. Table 1). Both the WT and NvClk1-/- populations exhibited a 24-hour periodicity in 12:12 hour Light-Dark (LD) cycles (Fig. 1d,e), with 15 out of 21 WT animals displaying 24-hour rhythmicity compared to only 9 out of 20 NvClk1-/- animals (Fig. 1f, Table 1, Supp. Table 2). The average acrophase for WT polyps (13.3 hours) was significantly lower than for NvClk1-/- polyps (17.3 hours) (Fig. 1f, Table 1).

Revised version: To assess the impact of the NvClkΔ mutation on circadian rhythm, we monitored the locomotor behavior of WT and NvClkΔ/Δ polyps under different light conditions (Supp. Table 1). Both the WT and NvClkΔ/Δ populations exhibited a 24-hour periodicity in 12:12 hour Light-Dark (LD) cycles (Fig. 1d, e), with 15 out of 21 WT animals displaying 24-hour rhythmicity compared to only 9 out of 20 NvClkΔ/Δ animals (Fig. 1f, Table 1, Supp. Table 2). The average acrophase for WT polyps (13.3 hours) was significantly earlier than for NvClkΔ/Δ polyps (17.3 hours) (Fig. 1f, Table 1). While we could detect a 24-hour rhythm for both genotypes, the delayed acrophase and reduced number of significant rhythmic polyps in the NvClkΔ/Δ suggest an alteration of the underlying rhythmicity mechanism.

Earlier version: We then investigated locomotor behavior under continuous conditions, namely continuous dark (DD) or continuous light (LL). WT polyps exhibited a 22-hour rhythmic behavior under both constant light conditions, with 17 out of 25 WT polyps displaying a 24-hour rhythm under DD and 7 out of 25 under LL (Fig. 1g-l). In contrast, few NvClk1-/- polyps displayed rhythmic behavior under constant conditions (1 out of 24 in DD and 1 out of 26 in LL) (Table 1). Additionally, we observed an intermediate phenotype in the locomotor behavior of heterozygous polyps for the NvClk1 allele in DD (Extended Data Fig. 1c-f).

Revised version: We then investigated locomotor behavior under continuous conditions, namely continuous dark (DD) or continuous light (LL). WT polyps exhibited a 22-hour rhythmic behavior under both constant light conditions, with 17 out of 25 WT polyps displaying a 24-hour rhythm under DD and 7 out of 25 under LL (Fig. 1g-l). In contrast, a few NvClkΔ/Δ polyps displayed rhythmic behavior under constant conditions (1 out of 24 in DD and 1 out of 26 in LL) (Table 1). Additionally, we observed an intermediate phenotype in the locomotor behavior of heterozygous polyps for the NvClkΔ allele in DD (Extended Data Fig. 1c-f). These results revealed that NvClkΔ/Δ polyps could not maintain circadian rhythmicity without diel light cues.

Earlier version: To further explore the 24-hour rhythm of NvClk1-/- polyps under LD, we tracked locomotor activity under a 6-hour light: 6-hour dark (LD 6:6) cycle after a regular diel 72-hour entrainment under 12:12 LD. While WT polyps maintained a marginally significant periodicity of 22 hours, NvClk1-/- polyps displayed a 12-hour rhythm at the population level (Fig. 1m-o). Specifically, we identified an important difference of 12- hour rhythmic individual polyps between WT and NvClk1-/- groups (1 out of 25 WT polyps vs. 13 out of 26 NvClk1-/- polyps) (Table 1). Notably, entrainment with LD 6:6 did not lead to a 12-hour rhythm in DD for both WT and NvClk1-/- polyps (Extended Data Fig. 1g-i).

Revised version: A 24-hour rhythm of NvClkΔ/Δ polyps under LD conditions could be attributed to either direct light response or the partial functioning of the circadian clock due to the nature of the mutation. To distinguish between these two possibilities, we monitored locomotor activity under a 6-hour light: 6-hour dark (LD 6:6) cycle after a regular diel 72-hour entrainment under 12:12 LD. While WT polyps maintained a marginally significant periodicity of 22 hours, NvClkΔ/Δ polyps displayed a 12-hour rhythm at the population level (Fig. 1m-o). Specifically, we identified a clear difference of 12-hour rhythmic individual polyps between WT and NvClkΔ/Δ groups (1 out of 25 WT polyps vs. 13 out of 26 NvClkΔ/Δ polyps) (Table 1). Notably, entrainment with LD 6:6 did not lead to a 12-hour rhythm in DD for both WT and NvClkΔ/Δ polyps (Extended Data Fig. 1g-i). These results support the hypothesis that the 24-hour rhythm observed in the NvClkΔ/Δ polyps in LD condition is due to the light-response pathway and not from an endogenous oscillator.

NvClk regulates rhythmic gene expression differentially in response to light conditions

Earlier version: To investigate the underlying molecular correlates of the behavioral phenotype found in NvClk1-/- polyps, we conducted transcriptional profiling. WT and NvClk1-/- polyps were sampled seven times every 4 hours over 24 hours under LD and DD conditions (Fig. 2a). We defined true rhythmic genes using a combination of both RAIN (rhythmicity analysis incorporating non-parametric methods) and JTK_CYCLE algorithms (Fig. 2b, Supp. Table 3). Within our dataset, we identified 119 genes rhythmic under LD and 107 rhythmic genes under DD in WT polyps (Fig. 2b). In NvClk1-/- polyps, we detected 147 rhythmic genes under LD and only 37 under DD (Fig. 2b).

Revised version: We conducted transcriptional profiling to investigate the underlying molecular correlates of the behavioral phenotype found in NvClkΔ/Δ polyps. WT and NvClkΔ/Δ polyps were sampled seven times every 4 hours over 24 hours under LD and DD conditions (Fig. 2a). To identify rhythmic genes, we employed stringent statistical parameters, including Benjamini-Hochberg (BH.Q) for the JTK method23 and adjusted p-value (p.adj) for the RAIN method24. This resulted in the identification of a minimal number of rhythmic genes. We detected only six rhythmic genes under LD conditions in WT polyps using the JTK method and 40 rhythmic genes using the RAIN method (Supp. Table 3). In DD condition, in the WT polyps, only two rhythmic genes were identified using the RAIN method (Supp. Table 3). Despite the risk of false positives, we opted not to use multiple testing but instead proposed to combine the JTK and RAIN algorithms to identify rhythmic genes, ensuring a robust approach to data analysis (p<0.01). We identified 119 rhythmic genes rhythmic under LD and 107 rhythmic genes under DD in WT polyps (Fig. 2b, Supp. Table 3). In NvClkΔ/Δ polyps, we detected 147 rhythmic genes under LD and only 37 under DD (Fig. 2b, Supp. Table 3).

Earlier version: NvClk1-/- shows rhythmic genes reduction in constant darkness with altered rhythmic features.

Revised version: NvClkΔ/Δ shows rhythmic gene reduction in constant darkness with altered rhythmic features.

Earlier version: (a) Overview of the experimental design used to generate RNA-seq data. Polyps were entrained for 72 hours before sampling at 4-hour intervals over a 24-hour period (dark arrows) in both LD and DD cycles. (b) Venn diagram comparing the total number of 24h rhythmic genes identified in WT and NvClk1-/- in LD and DD cycles with a p <0.01 with RAIN and JTK. (c) Average acrophase comparison between rhythmic genes in LD and DD in WT polyps. Mann-Whitney test, p<0.001. (d) Average acrophase comparison between rhythmic genes in LD and DD in NvClk1- polyps. Mann-Whitney test, p>0.05. (3) Average relative amplitude comparison between rhythmic genes in LD and DD in WT polyps. Mann-Whitney test, p<0.05. (d) Average relative amplitude comparison between rhythmic genes in LD and DD in NvClk1-/- polyps. Mann-Whitney test, p>0.05. (c-f) sample size (n) indicated below each boxplots.

Revised version: (a) Overview of the experimental design used to generate RNA-seq data. Polyps were entrained for 72 hours before sampling at 4-hour intervals over a 24-hour period (dark arrows) in both LD and DD cycles. (b) Venn diagram comparing the total number of 24h rhythmic genes identified in WT and NvClkΔ/Δ in LD and DD cycles with a p <0.01 with RAIN and JTK. (c) Average acrophase comparison between rhythmic genes in LD and DD in WT polyps. Mann-Whitney test, p<0.001. (d) Average acrophase comparison between rhythmic genes in LD and DD in NvClkΔ/Δ polyps. Mann-Whitney test, p>0.05. (e) Average relative amplitude comparison between rhythmic genes in LD and DD in WT polyps. Mann-Whitney test, p<0.05. (d) Average relative amplitude comparison between rhythmic genes in LD and DD in NvClkΔ/Δ polyps. Mann-Whitney test, p>0.05. (c-f) sample size (n) indicated below each boxplot.

Earlier version: The rhythmic genes in WT polyps displayed a higher acrophase under DD compared to LD (17.20h vs 12.93h, Fig. 2c). However, no differences were detected between LD and DD rhythmic genes in NvClk1-/- polyps (Fig. 2d). Similarly, the relative amplitude (rAMP) of DD rhythmic genes was higher in WT polyps compared to LD (0.61 vs 0.43, Fig. 2e), but no rAMP difference was observed between LD and DD rhythmic genes in NvClk1-/- polyps (Fig. 2f).

Revised version: The rhythmic genes in WT polyps displayed a delayed acrophase under DD compared to LD (17.20h vs. 12.93h, Fig. 2c). However, no differences were detected between LD and DD rhythmic genes in NvClkΔ/Δ polyps (Fig. 2d). Similarly, the relative amplitude (the gene amplitude divided by its baseline, rAMP) of DD rhythmic genes was higher in WT polyps compared to LD (0.61 vs. 0.43, Fig. 2e), but no rAMP difference was observed between LD and DD rhythmic genes in NvClkΔ/Δ polyps (Fig. 2f).

Earlier version: We performed a clustering analysis on the rhythmic genes using the DPGP model (Dirichlet process Gaussian process mixture model). The average number of genes per cluster between LD and DD conditions in WT polyps did not differ significantly (7.3 vs 7.6, Fig. 3a, Supp. Table 4). However, in NvClk1-/- polyps, the average number of genes per cluster was significantly lower in DD compared to LD (4.1 vs 8.6, Fig. 3b). Specifically, in WT DD polyps, we observed clusters with high numbers of genes peaking at subjective night (Fig. 3c, Supp. Table 4). We did not identify any GO term enrichment to any of the clusters.

Revised version: Are rhythmic genes organized into “transcriptional time clusters”? Does the NvClkΔ mutation modify cluster recruitments, causing the loss of rhythmic behavior under DD conditions? We performed a clustering analysis on the rhythmic genes using the DPGP model (Dirichlet process Gaussian process mixture model). The number of genes per cluster between LD and DD conditions in WT polyps did not differ significantly (7.3 vs. 7.6, Extended Data Fig. 2a, Supp. Table 4). Interestingly, when clusters are organized by their acrophase, we observed clusters with higher numbers of genes peaking at subjective night in WT under DD conditions (Extended Data Fig. 2b, Supp. Table 4). In NvClkΔ/Δ polyps, the number of genes per cluster was significantly reduced in DD compared to the LD condition (4.1 vs. 8.6, Extended Data Fig. 2c,d). We did not identify GO-term enrichment in any cluster. However, the overlap between clusters and behavior opens new directions for further functional analysis (Extended Data Fig. 3b,d and Supp. Table 4). Overall, the reduced number of rhythmic genes in NvClkΔ/Δ polyps under the DD condition and the reduced number of genes per cluster confirm the necessity of NvClk to recruit rhythmic genes in the DD condition and to organize them in transcriptional time clusters.

Removed in the revised version: NvClk1-/- shows temporal gene clusters recruitment alteration in constant darkness.

Removed in the revised version: (a) Gene per cluster count comparison between LD and DD in WT polyps. Number of cluster indicated below each boxplot. Mann-Whitney, p > 0.05. (b) Gene per cluster comparison between LD and DD in NvClk1-/-. Sample size (n) indicated below each boxplot. Mann-Whitney, p < 0.05. (c) Representative temporal gene cluster for each condition. For both light conditions, from top to bottom cluster peaking at day, then twilight, then at night. Number of gene for each specific cluster indicated on the x-axis.

Earlier version: NvClk regulates temporal expression pattern of pacemaker genes

Revised version: NvClk regulates the temporal expression pattern of pacemaker genes

Added in the revised version: In line with previous findings in Nematostella12,14, candidate pacemaker genes showed arrhythmic expression under DD conditions (Fig. 3a, Supp. Table 3). However, the altered expression patterns observed in NvClkΔ/Δ compared to WT polyps in LD condition showed increased transcripts for some genes (i.e., NvClk and NvPar-bzipd). In contrast, others (NvCipc and NvPar-bzipc) exhibited a reduction in transcript numbers (Fig. 3a, Supp. Table 3). If we hypothesize that the first two genes (NvClk and NvPar-bzipd) act as positive factors and the latter two (NvCipc and NvPar-bzipc) potentially serve as negative regulator of the former, the lack of functionality of the NvClkΔ allele would explain the observed difference in transcript levels between NvClkΔ/Δ and WT polyps.

Removed in the revised version: In line with previous findings in Nematostella12,14, candidate pacemaker genes showed arrhythmic expression under DD conditions (Fig. 4a, Supp. Table 3). However, the altered expression patterns observed in LD NvClk1-/- polyps, compared to LD WT polyps (Fig. 4a, Supp. Table 3), suggest a complex regulatory effect.

Earlier version: NvClk1-/- alters temporal pacemaker genes expression.

Revised version: NvClkΔ/Δ alters temporal pacemaker gene expression.

Earlier version: (a) Four pacemaker gene are individually plotted showing the read count over 24h in LD and DD in WT (black) and NvClk1-/- (red). The continuous line represents significant rhythmicity (RAIN&JTK p<0.01) while the dashed line indicates no rhythmicity. (b) correlation matrix of candidate pacemaker genes expression in LD for WT on the left and NvClk1-/- on the right. (c) schematic representation of the promoter sequences analyses 5kb upstream the putative ATG. Black boxes represent canonical E-boxes while circadian E-boxes are green. Below the logo motif we used to identify canonical and circadian Ebox. (d) Circadian / Canonical ratio (in %) per condition. Kruskal-Wallis, multiple comparison, a vs b : p<0.05.

Revised version: (a) Four pacemaker genes are plotted, showing the read counts over 24h in LD and DD in WT (black) and NvClkΔ/Δ (red). The continuous line represents significant rhythmicity (RAIN&JTK p<0.01), while the dashed line indicates no rhythmicity. (b) correlation matrix of candidate pacemaker genes expression in LD for WT on the left and NvClkΔ/Δ on the right. (c) schematic representation of the promoter sequences analyses 5kb upstream of the putative ATG. Black boxes represent canonical E-boxes, while circadian E-boxes are green. Below is the logo motif we used to identify canonical and circadian Ebox. (d) Circadian / Canonical ratio (in %) per condition. Kruskal-Wallis, multiple comparison, a vs b : p<0.05.

Added in the revised version: To systematically assess the mutation’s impact on all the potential pacemaker genes, we utilized a correlation matrix based on their temporal transcript number levels, offering a comprehensive overview of their temporal organization. In WT polyps under LD conditions, the clustering categorized genes into two groups: one exhibiting a daytime peaking, containing NvClk, and another peaking at night comprising NvPar-bzipc and NvCipc. Notably, in LD NvClkΔ/Δ polyps, this second cluster contained two additional genes and displayed a weakened anticorrelation with the NvClk cluster (Fig. 3b). These observations suggest that the pacemaker oscillation, generated by the interplay of positive and negative feedback loops, relies on the precise temporal organization of these potential pacemaker factors into distinct clusters. The disruption of this organization by the NvClkΔ allele underscores the central role of NvClk in pacemaker function.

Removed in the revised version: To systematically analyze the mutation’s impact on these genes, we constructed a correlation matrix based on their expression levels. In LD WT polyps, genes clustered into two groups: one with NvClk peaking during the day, anticorrelated to the second cluster (NvPar-bzipc and NvCipc) peaking at night. In LD NvClk1-/- polyps, the second cluster contained two additional genes and showed weaker anticorrelation with the NvClk cluster (Fig. 4b).

Earlier version: To go further into the regulatory mechanisms downstream of the pacemaker, we examined the presence of circadian E-box motifs (CACGTG) within 5kb upstream of the predicted ATG of rhythmic genes. We calculated circadian/canonical E-box enrichment to account for the total variation in the number of canonical E-boxes (Fig. 4c). Notably, only the candidate pacemaker genes exhibited a significant enrichment in circadian E-boxes in their promoters (15.9%) compared to the WT (5.6%), NvClk1-/- (4.8%) rhythmic genes, and non-rhythmic genes (6.8%) (Fig. 4d).

Revised version: To go further into the regulatory mechanisms downstream of the pacemaker, we examined the presence of circadian E-box motifs (CACGTG) within 5kb upstream of the predicted ATG of rhythmic genes. We calculated circadian/canonical E-box enrichment to account for the total variation in the number of canonical E-boxes (Fig. 3c). Notably, only the candidate pacemaker genes exhibited a significant enrichment in circadian E-boxes in their promoters (15.9%) compared to the WT (5.6%), NvClkΔ/Δ (4.8%) rhythmic genes, and non-rhythmic genes (6.8%) (Fig. 3d).

NvClk coordinates cell division and neuronal pathways in constant darkness

Earlier version: In addition to the transcriptomic rhythmic analysis, our aim was to identify processes regulated by NvClk that may not necessarily exhibit rhythmicity. To achieve this, we conducted a differential gene expression analysis on the total transcriptome between genotypes under each light condition. Under LD conditions, NvClk1-/- polyps exhibited 457 down-regulated genes and 646 up-regulated genes, with no significant enrichment in GO terms observed (Fig. 5a, Supp. Table 4 and 5). However, in DD conditions, NvClk1-/- displayed 2450 down-regulated genes and 1770 up-regulated genes (Fig. 5b, Supp. Table 4). Notably, the down-regulated genes were predominantly enriched in processes related to mitosis, microtubules, and ciliary/flagellar motility. Conversely, the up-regulated genes showed significant enrichment in processes such as the modulation of another organism’s processes, axonal guidance, and sensory perception (Fig. 5b, Supp. Table 5).

Revised version: In addition to the transcriptomic rhythmic analysis, we aimed to identify processes regulated by NvClk that may not necessarily exhibit rhythmicity. We conducted a differential gene expression analysis on the total transcriptome between genotypes under each light condition to achieve this. Under LD conditions, NvClkΔ/Δ polyps exhibited 457 down-regulated genes and 646 up-regulated genes, with no significant enrichment in GO terms observed (Fig. 4a, Supp. Table 4 and 5). However, in DD conditions, NvClkΔ/Δ displayed 2450 down-regulated genes and 1770 up-regulated genes (Fig. 4b, Supp. Table 4). Notably, we identified enrichment in down-regulated genes in processes related to mitosis, microtubules, and ciliary/flagellar motility. Conversely, the up-regulated genes showed significant enrichment in processes such as the modulation of another organism’s processes, axonal guidance, and sensory perception (Fig. 4b, Supp. Table 5).

Earlier version: NvClk1-/- disrupts cell-cycle and neuronal pathways in constant darkness.

Revised version: NvClkΔ/Δ disrupts cell-cycle and neuronal pathways in constant darkness.

Earlier version: (a) Volcano plot showing the differential expression of genes (DEG) between WT and NvClk1-/- in LD (Left) and DD (right). Dashed line indicate the threshold used to detect DEG (p.adj<0.01). Red dots indicate down regulated genes and black dots up-regulated genes in NvClk1-/- compare to WT polyps (b) Gene Ontology (GO) terms with with significant fold-enrichment (Bonferroni corrected p-value or p.adjusted <0.01) for the DEG analysis in DD. Down regulated genes in Red while Up regulated genes in Black.

Revised version: (a) Volcano plot showing the differential expression of genes (DEG) between WT and NvClkΔ/Δ in LD (Left) and DD (right). Dashed line indicates the threshold used to detect DEG (p.adj<0.01). Red dots indicate down regulated genes and black dots up-regulated genes in NvClkΔ/Δ compare to WT polyps (b) Gene Ontology (GO) terms with with significant fold-enrichment (Bonferroni corrected p-value or p.adjusted <0.01) for the DEG analysis in DD. Down regulated genes in Red while Up regulated genes in Black.

Discussion

Conserved behavioral CLOCK function through animal evolution

Earlier version: Our study provides valuable insights into the evolution of circadian clocks by characterizing the effects of the first Clock mutation in a cnidarian, the sea-anemone Nematostella vectensis. Through our behavioral assays, we have shown that NvClk is essential for maintaining rhythmic locomotor activity in the absence of an entraining light cue. Although the rhythmicity of the NvClk1+/- heterozygote polyps was affected in DD, our results could not discriminate a dominant-negative from a full loss of function to identify the nature of this mutation (Extended Data Fig. 1g-i). Studies in various model organisms further support the importance of CLOCK in regulating circadian locomotion. For instance, homozygous DmClk-/- flies exhibit a loss of circadian locomotion in constant darkness23,24, while the DnClk1a zebrafish mutant displayed a shortened period under the same conditions25. Interestingly, MmClock-/- mice maintain robust circadian locomotion in constant darkness, suggesting potential functional redundancy among mammalian clock genes26–29. Overall, these findings support a conserved role of CLOCK in preserving circadian behavioral rhythms in absence of light cues across the distant Nematostella, flies, zebrafish, and mice.

Revised version: Our study provides valuable insights into the evolution of circadian clocks by characterizing the effects of the first Clock mutation in a cnidarian, the sea anemone Nematostella vectensis. Our behavioral assays showed that NvClk is essential for maintaining rhythmic locomotor activity without an entraining light cue. Although the rhythmicity of the NvClk+/Δ heterozygote polyps was affected in DD, our results could not discriminate a dominant-negative from a total loss of function to identify the nature of this mutation (Extended Data Fig. 1g-i). Studies in various model organisms further support the importance of CLOCK in regulating circadian locomotion. For instance, DmClkar/ar flies exhibit a loss of circadian locomotion in constant darkness25, while the DnClk1adg3/dg3 zebrafish mutant displayed a shortened period under the same conditions26. The dominant-negative mutant MmClock Δ5–6/Δ5–6 mice showed a loss of circadian locomotion in constant darkness, however the complete deletion of the MmClock gene did not affect the circadian behavior rhythm in constant darkness suggesting compensation by a paralog27–29. Overall, these findings support a conserved role of CLOCK in preserving circadian behavioral rhythms in the absence of light cues across the distant Nematostella, flies, zebrafish, and mice.

Earlier version: Moreover, the conservation of a 24-hour locomotion rhythm in LD of the NvClk1-/- polyps with a different mean acrophase, revealed a light-response pathway independent of the circadian circuit, consistent with observations in other animal models23,26,30 (Fig. 1f). NvClk1-/- polyps exposed to a 12:12h LD cycle exhibited a 24-hour period, whereas those exposed to a 6:6h LD cycle displayed a 12-hour period. Notably, nearly no WT polyps exhibited a 12-hour rhythm under this condition, suggesting that the circadian clock overrides the light-response pathway (Fig. 6a). While some of the circadian factors can directly sense the light such as CRY proteins31, 29 putative NvOpsin have been identified in the genome which could be involved in the light-response pathway32. Behavioral tracking of NvClk1-/- polyps exposed to different wavelength could help to identify candidates for further functional studies of the light-response pathway.

Revised version: Moreover, the conservation of a 24-hour locomotion rhythm in LD of the NvClkΔ/Δ polyps with a delayed acrophase revealed a light-response pathway independent of the circadian circuit, consistent with observations in other animal models25,27,30 (Fig. 1f). NvClkΔ/Δ polyps exposed to a 12:12h LD cycle exhibited a 24-hour period. In contrast, those exposed to a 6:6h LD cycle displayed a 12-hour period. Notably, nearly no WT polyps exhibited a 12-hour rhythm under this condition, suggesting that the circadian clock overrides the light-response pathway (Fig. 5a). While some of the circadian factors can directly sense the light, such as CRY proteins31, 29 putative NvOpsin have been identified in the genome which could be involved in the light-response pathway32. Behavioral tracking of NvClkΔ/Δ polyps exposed to different wavelengths could help to identify candidates for further functional studies of the light-response pathway.

Summary of NvClk function in the regulation of Nematostella circadian rhythmicity.

Transcriptional rhythmicity plasticity downstream NvClk

Earlier version: At the transcriptomic level, previous studies in Nematostella have shown large changes in the transcriptional profile of many genes after a single day of constant darkness, including the candidate pacemaker genes that were found arrhythmic, despite sustain circadian locomotion12,14,22. Consistent with previous transcriptomic analysis in cnidarian7,14,16,33, most of rhythmic genes identified in LD were different from those identified in DD in the WT polyps. Notably, they displayed higher mean acrophase and larger mean amplitude in DD suggesting a differential regulation in response to light conditions, which has not been investigated in previous cnidarian studies. Additionally, an overlap between our LD rhythmic genes and Leach et al. dataset14 suggests the robustness of pacemaker rhythmic transcription in LD conditions compared to downstream rhythmic genes (Extended data Fig. 2). Overall, these observations suggest a potential plasticity of the pacemaker in selecting specific genes depending on environmental conditions.

Revised version: At the transcriptomic level, previous studies in Nematostella have shown large changes in the transcriptional profile of many genes after a single day of constant darkness, including the candidate pacemaker genes that were found arrhythmic despite sustaining circadian locomotion12,14,22. Consistent with previous transcriptomic analysis in cnidarian7,14,16,33, most of the rhythmic genes identified in LD differed from those identified in DD in the WT polyps. Notably, they displayed higher mean acrophase and larger mean amplitude in DD, suggesting a differential regulation in response to light conditions, which had not been investigated in previous cnidarian studies. Furthermore, the overlap observed between our LD rhythmic genes and those identified by Leach et al.14 underscores the robustness of pacemaker rhythmic transcription in LD conditions (Extended data Fig. 3). However, the lack of overlap for rhythmic genes raises intriguing questions. Differences in experimental conditions, including genetic backgrounds, light system (Neon vs. LED), salinity (12ppt vs. 15ppt), and temperature (17°C vs. 25°C), may contribute to these discrepancies. Further investigations are necessary to determine if the lack of overlap of rhythmic genes downstream of the potential pacemaker genes results from an organism’s adaptation to its environment and, therefore, reflects the plasticity of the pacemaker in regulating its downstream rhythmic genes.

Added in the revised version: Our study identified 24-hour rhythmic behavior in NvClkΔ/Δ polyps under LD conditions, suggesting an alternative mechanism for generating molecular rhythmicity via the light-sensing pathway. However, it is crucial to note minimal overlap between the rhythmic genes identified in NvClkΔ/Δ and WT polyps under LD conditions. This discrepancy indicates that the light-response pathway may not fully replicate the normal pacemaker functions observed in WT polyps, highlighting the need for further investigation into the recruitment and function of these genes. Additionally, the reduced number of rhythmic genes identified in NvClkΔ/Δ polyps under the DD condition underscores the crucial role of NvClk in maintaining molecular rhythm without light cues.

Removed in the revised version: Similar to the rhythmic behavior of NvClk1-/- polyps under LD conditions, the identification of 24-hour rhythmic genes suggests potential compensation by the light-response pathway. However, most of rhythmic genes in the NvClk1-/- compared to WT polyps in LD conditions are different. This suggests the light-response pathway may not fully reproduce the normal circadian pattern observed in WT polyps and would require further investigation to understand their recruitment and function. Furthermore, the identification of a reduced number of rhythmic genes in DD NvClk1-/- polyps revealed the importance of NvClk in maintaining molecular rhythm in absence of light. However, the persistence of these rhythmic genes raises questions regarding their origin. They could be the result of the mutation itself, a compensatory mechanism or false positives resulting from our analysis. The study by Ray et al., which report controversial findings regarding RNA rhythmicity in the Bmal1-/- mouse cells (NvCycle orthologous in Nematostella) is an interesting case of rhythmicity analysis to consider34–38. Our attempt to identify rhythmic genes faced challenges, particularly in the choice of statistical parameters. Initially, using stringent criteria like BHQ for the JTK method or p.adj with RAIN resulted in a very limited number of genes for analysis (Supp. Table 3). Consequently, we opted for a balanced approach combining both algorithms with a p<0.01 threshold. We made a compromise to identify rhythmic genes with greater confidence, but we are aware that the methodological choices are critical for the results.

Earlier version: The clustering analysis revealed that rhythmic genes can be categorized by expression pattern similarity. This approach defined “temporal gene clusters” as group of seven/eight genes in average in the WT (Fig. 6b). Their discovery poses a fundamental question: How a group of genes are co-regulated in times and space (cell types) by the pacemaker? Their recruitment is disrupted in the DD NvClk1-/- polyps suggesting an essential function of NvClk in absence of light. The combination of published scAtlas39 and multiplexed FISH techniques40 will be essential to further investigate the biological regulation and function of these temporal gene clusters.

Revised version: The clustering analysis revealed that rhythmic genes can be categorized into “transcriptional time clusters” (aka synexpression clusters)16,34 by group of seven/eight genes in average in the WT (Fig. 5b). Their existence raises a fundamental question that has yet to be answered: How is a group of genes co-regulated in time and space (cell types) by the pacemaker? Their recruitment is disrupted in the DD NvClkΔ/Δ polyps suggesting an essential function of NvClk in absence of light. The combination of published scAtlas35 and multiplexed FISH techniques36 will be essential to further investigate the biological regulation and function of these transcriptional time clusters.

Earlier version: NvClk temporally organize pacemaker genes expression to drive rhythmic genes recruitment

Revised version: NvClk temporally organizes pacemaker gene expression to drive rhythmic gene recruitment

Earlier version: Our study reveals that NvClk plays a key role in regulating the temporal transcription of pacemaker candidate genes (Fig.4a). Our analysis identified two clusters of pacemaker genes: One containing NvClk, and a second one containing a potential NvClk inhibitor (NvCipc)41,42. The alteration of clusters composition with a weaker anticorrelation in LD NvClk1-/- polyps might generate a desynchronization of the pacemaker factors availability. Indeed, regulation of rhythmic transcription involved a complex protein-protein-DNA timing interaction. Furthermore, we did not identify any circadian E-boxes enrichment in rhythmic genes between conditions, except for the candidate pacemaker genes. Altogether, this supports the function of NvClk in orchestrating the timing interaction of pacemaker factors to select downstream rhythmic genes, indicating a more complex regulatory landscape at play.

Revised version: Our study reveals that NvClk plays a crucial role in regulating the temporal transcription of pacemaker candidate genes (Fig.3a). Our analysis identified two clusters of pacemaker genes: One containing NvClk and a second one containing a potential NvClk inhibitor (NvCipc)37,38. These two clusters suggest the organization of the potential pacemaker genes transcription into interlock feedback loops with antiphase peaks, probably at the origin of the pacemaker oscillator function21,39,40. The alteration of cluster composition with a weaker anticorrelation in LD NvClkΔ/Δ polyps might generate a desynchronization of the pacemaker factors’ availability. Indeed, regulation of rhythmic transcription involved a complex protein-protein-DNA timing interaction. Furthermore, we did not identify any circadian E-boxes enrichment in rhythmic genes between conditions, except for the candidate pacemaker genes. Altogether, this supports the function of NvClk in orchestrating the timing interaction of pacemaker factors to select downstream rhythmic genes, indicating a more complex regulatory landscape at play.

Earlier version: However, one significant unanswered question in our study is the reason for the arrhythmic transcription of putative pacemaker genes in DD. We hypothesize that using whole animals for sampling material might mask oscillating gene expression signals, especially if signals are present in a small number of cells or if tissues exhibit rhythmic gene expression in different phases. Furthermore, it is important to acknowledge a limitation in our interpretation, common in chronobiology: the use of RNA oscillation as a proxy for protein oscillation and function. The development of tools to study the pacemaker factors at the protein level in Nematostella will leverage this limitation in the field.

Revised version: However, one significant unanswered question in our study is the reason for the arrhythmic transcription of putative pacemaker genes in DD. Using whole animals for sampling material might mask oscillating gene expression signals, especially if signals are present in a small number of cells or if tissues exhibit rhythmic gene expression in different phases. Furthermore, we must acknowledge a limitation in our interpretation, which is common in chronobiology: using RNA oscillation as a proxy for protein oscillation and function. The development of tools to study the pacemaker factors at the protein level in Nematostella will leverage this limitation in the field.

Earlier version: NvClk regulates processes involved in cell proliferation and the neural system in absence of light

Revised version: NvClk regulates processes involved in cell proliferation and the neural system in the absence of light

Earlier version: Our study of NvClk suggests coordination of cellular processes, especially in the absence of light. The results of our rhythmic transcriptomic analysis (Fig. 2-3) raised questions regarding indirect effects, and the non-rhythmic function of NvClk. We performed a differential gene expression analysis on the total transcriptome for each light condition. Under LD conditions, NvClk1-/-polyps exhibited significant changes in gene expression but no GO term enrichment was found (Fig. 5a, Supp. Table 5) revealing multiple altered processes we cannot yet identified. In contrast, under DD conditions, NvClk1-/- polyps displayed more pronounced alterations, with more DEGs and enriched GO-terms for down-regulated genes related to mitosis, microtubule organization, and ciliary/flagellar motility, while the up-regulated genes showed enrichment in processes such as the modulation of other organism’s processes, axonal guidance, and sensory perception (Fig. 5b, Supp. Table 5). These results imply that NvClk has non-circadian functions that are dependent on light availability. This is particularly noteworthy considering the expression of core pacemaker genes, which are known to be arrhythmic during larvae stages, potentially involved in developmental processes22.

Revised version: Our study of NvClk suggests coordination of cellular processes, especially in the absence of light. Our rhythmic transcriptomic analysis results (Fig. 2 and Extended data Fig. 2) raised questions regarding indirect effects and the non-rhythmic function of NvClk. We performed a differential gene expression analysis on the total transcriptome for each light condition. Under LD conditions, while NvClkΔ/Δ polyps exhibited significant changes in gene expression, we could not identify any GO term enrichment (Fig. 4a, Supp. Table 5), revealing multiple altered processes we cannot yet identify.

Added in the revised version: In contrast, under DD conditions, NvClkΔ/Δ polyps displayed more pronounced alterations, with more DEGs and enriched GO-terms for down-regulated genes related to mitosis, microtubule organization, and ciliary/flagellar motility, while the up-regulated genes showed enrichment in processes such as the modulation of other organism’s processes, axonal guidance, and sensory perception (Fig. 4b, Supp. Table 5). These results imply that NvClk has non-circadian functions dependent on light availability. This is particularly noteworthy considering the expression of core pacemaker genes, known to be arrhythmic during larvae stages, potentially involved in developmental processes22.

Earlier version: Taken as a whole, this study provides novel insights into circadian regulation in Nematostella vectensis and sheds light on the evolutionary origin of circadian time maintenance. Our findings indicate that CLOCK function is conserved from Cnidaria to mammals to maintain rhythmicity in absence of light cues. Furthermore, it revealed a light-response pathway able to compensate at both behavioral and molecular level using light cues. This circadian clock mutant opens new avenues for investigating cell-type-specific mechanisms of the circadian clock that drive the molecular and phenotypical oscillations of cnidarians. By further exploring the circadian clock mechanisms in cnidarians, we can gain deeper insights into the evolutionary origins of this critical aspect of biology, enhancing our understanding of how organisms have evolved to keep track of time and adapt to their environment.

Revised version: This study provides novel insights into circadian regulation in Nematostella vectensis and sheds light on the evolutionary origin of circadian time maintenance. Our findings indicate that CLOCK function is conserved from cnidaria to mammals to maintain rhythmicity without diel light cues. Furthermore, it revealed a light-response pathway able to compensate at both behavioral and molecular levels using light cues. This circadian clock mutant opens new avenues for investigating cell-type-specific mechanisms of the circadian clock that drive the molecular and phenotypical oscillations of cnidarians. By further exploring the circadian clock mechanisms in cnidarians, we can gain deeper insights into the evolutionary origins of this critical aspect of biology, enhancing our understanding of how organisms have evolved to keep track of time and adapt to their environment.

Supporting information